Glycosylated IgG antibodies contribute to the recovery of haemorrhagic fever with renal syndrome patients

Author response:

The following is the authors’ response to the previous reviews

Reviewers 1:

Summary:

The authors investigated the potential role of IgG N-glycosylation in Haemorrhagic Fever with Renal Syndrome (HFRS), which may offer significant insights for understanding molecular mechanisms and for the development of therapeutic strategies for this infectious disease.

While the majority of the issues have been addressed, a few minor points still remain unresolved. Quality control should be conducted prior to the analysis of clinical samples. However, the coefficient of variation (CV) value was not provided for the paired acute and convalescent-phase samples from 65 confirmed HFRS patients, which were analyzed to assess inter-individual biological variability. It is important to note that biological replication should be evaluated using general samples, such as standard serum.

We thank the reviewer for this insightful and critical comment regarding the quality control of our analytical data and the assessment of biological variability. We agree that this is essential for validating the reliability of our findings. We have now provided the requested CV data and clarified this point in the revised manuscript as detailed below.

"This dual-replicate strategy enabled a comprehensive evaluation of both biological heterogeneity and assay precision, and the coefficient of variation for samples were below 16%." Please see the Materials and Methods (Page 16, lines 360-362, and Author response table 1).

Reviewers 2:

This work sought to explore antibody responses in the context of hemorrhagic fever with renal syndrome (HFRS) - a severe disease caused by Hantaan virus infection. Little is known about the characteristics or functional relevance of IgG Fc glycosylation in HFRS. To address this gap, the authors analyzed samples from 65 patients with HFRS spanning the acute and convalescent phases of disease via IgG Fc glycan analysis, scRNAseq, and flow cytometry. The authors observed changes in Fc glycosylation (increased fucosylation and decreased bisection) coinciding with a 4-fold or greater increased in Haantan virus-specific antibody titer. The study also includes exploratory analyses linking IgG glycan profiles to glycosylation-related gene expression in distinct B cell subsets, using single-cell transcriptomics. Overall, this is an interesting study that combines serological profiling with transcriptomic data to shed light on humoral immune responses in an underexplored infectious disease. The integration of Fc glycosylation data with single-cell transcriptomic data is a strength.The authors have addressed the major concerns from the initial review. However, one point to emphasize is that the data are correlative. While the associations between Fc glycosylation changes and recovery are intriguing, the evidence does not establish causation. This is not a weakness, as correlative studies can still be highly valuable and informative. However, the manuscript would be strengthened by making this distinction clear, particularly in the title.

The verb "accelerated" in the title implies that the glycosylation state of IgG was a direct driver of recovery, rather than something that correlated with recovery. Thus, a more neutral word/phrase would be ideal.

We sincerely thank the reviewer for this insightful suggestion. We agree that the use of "accelerated" might overstate the potential role of IgG glycosylation, which has not been clearly clarified by our current findings. As reported in results (particularly in Figure 2), partial glycosylation exhibits statistically significant variations between seropositive and seronegative statuses, before and after seroconversion, and across different HTNV- NP specific antibody titers. Therefore, we have replaced "accelerated" with "contribute to" in the Title: "Glycosylated IgG antibodies contribute to the recovery of haemorrhagic fever with renal syndrome patients".

Changes from the earlier version to the revised version

Abstract

Haemorrhagic fever with renal syndrome (HFRS) is a fatal disease caused by Hantaan virus (HTNV) infection. Humoral immunity is essential for effective viral clearance, however, the glycosylation characteristics of immunoglobulin G (IgG) in HFRS patients is not well known. Peripheral blood mononuclear cells from HFRS patients were obtained for B subset analysis using scRNA-seq and flow cytometry. HTNV specific IgG antibody titers were detected by enzyme-linked immunosorbent assay, and IgG glycosylation was analyzed by ultra-performance liquid chromatography. The proportions of the antibody-secreting memory (ASM) B cells and plasmablasts (PB) were significantly expanded among acute HFRS patients. We discovered significantly increased fucosylated IgG and decreased bisecting N-acetylglucosamine during the convalescent phase of HTNV infection. Meanwhile, positive correlations were observed between ASM subsets and galactosylation/sialylation in the IgG Fc region, and between PB subsets and sialylation. Notably, the glycosylation-related genes, such as RPN1 and RPN2, were primarily expressed differentially in the ASM and PB subclusters which were enriched in the N-glycosylation modifications of proteins through asparagine. Our findings indicated that IgG N-glycosylation may play a crucial role in combating HTNV infection and contributing to clinical recovery, which provided new insights for optimizing glycoengineered therapeutic antibodies.

Introduction

The causative agent of haemorrhagic fever with renal syndrome (HFRS) is Hantaan virus, which belongs to the family Hantaviridae and is primarily transmitted by rodents (Zheng et al., 2019). China is the most heavily affected country by HFRS, accounting for >90% of global cases, and a total of 32,462 cases of HFRS were reported in China during 2019–2022, with a fatality rate ranging from 0.1 to 15% (R. Li et al., 2023) (https://www.chinacdc.cn).

Hantaviruses are tri-segmented, single-stranded, negative-sense RNA viruses, whose genomes consist of three regions: large (L), medium (M), and small (S). The glycoproteins Gn and Gc, encoded by the M segment, can infect target cells - primarily vascular endothelial cells - via β3 integrin receptors (Pizarro et al., 2019). Simultaneously, they could also infect other cell types, such as mononuclear macrophages and dendritic cells, leading to systemic viral infection. Although hantavirus replication is thought to occur primarily in the vascular endothelium without direct cytopathic effects, a plethora of innate immune cells mediate host antiviral defenses. These include natural killer cells, neutrophils, monocytes, and macrophages, together with pattern recognition receptors (PRRs), interferons (IFNs), antiviral proteins, and complement activation, e.g., via the pentraxin 3 (PTX3) pathway, which can exacerbate HFRS disease progression leading to immunopathological damage through cytokine/chemokine production, cytoskeletal rearrangements in endothelial cells, ultimately amplifying vascular dysfunction (Tariq & Kim, 2022). Rapid and effective humoral immune responses, however, such as neutralizing antibody responses targeting the glycoproteins Gn/Gc, contribute to rapid recovery from HFRS and are critical for protection from severe disease (Engdahl & Crowe, 2020; Li et al., 2020).

Immunoglobulin G (IgG) N-linked glycosylation mediates critical functions modulating antiviral immunity during viral infection. Changes in the conserved N-linked glycan Asn297 in the Fc region of IgG typically by fucosylation, galactosylation, or sialylation can alter antibody effector function. A reduction in core fucosylation decreases IgG binding to NK cell FcγRIIIa promotes antibody-dependent cellular cytotoxicity (ADCC) necessary for clearance of viruses, including SARS-CoV-2, dengue and HIV-1 whereas sialylation can attenuate immune responses resulting in immune evasion (Ash et al., 2022; Haslund-Gourley et al., 2024; Hou et al., 2021; Wang et al., 2017). Changes in IgG and other protein N-linked glycosylation profiles therefore shape virus-host interactions and disease progression.

Importantly, there have not been prior studies specifically examining plasma IgG N-glycome profiles derived from chromatographic peak data in HFRS patients, particularly in relation to seroconversion status. This gap in our knowledge motivated our systematic investigation of both total and virus-specific IgG glycosylation dynamics during acute infection. In the present study, we combined scRNA-seq and flow cytometry to reveal the phenotypes of B cell responses during Hantaan virus (HTNV) infection, and explored the transcriptomic features of reactive B cell subsets. A total of 166 HFRS patients, including 65 paired HFRS samples, were employed to profile the IgG-Fc glycosylation pattern and to investigate the potential regulatory role of IgG-Fc glycosylation in HFRS pathogenesis.

Results

B cell compositional characteristics in HFRS patients

To characterize the humoral immune profiles in HFRS patients, we enrolled 166 suspected HTNV-infected patients who were admitted to Baoji Central Hospital in Shaanxi Province, China, between October 2019 and January 2022. Among them, 65 met the inclusion criteria and were included in the study (Figure 1). We identified a total of eight cell subpopulations, including subsets of CD4 T cells, CD8 T cells, CD14 monocytes, NK cells, B cells, platelets, endothelial progenitor cells, and red blood cells (Figure 2—figure supplement 1A and B). The proportions of these cell types were comparable between acute HFRS patients and healthy controls, with no significant batch effects observed (Figure 2—figure supplement 1C). CD4 T cells (31.4%) were the predominant subset prior to HTNV infection, whereas CD8 T cells (44%) became the dominant population post-infection. Notably, the proportions of NK cells (15.9% to 5.6%, p<0.001) and platelets (3.1% to 0.6%, p<0.01) significantly decreased, whereas the level of CD14 monocytes (15.9% to 32.3%, p<0.05) nearly doubled. The overall proportions of B cells, however, did not exhibit significant changes (Figure 2—figure supplement 1D and E).

Flow chart of HFRS patient recruitment in Baoji Central Hospital in Shaanxi province from November 2019 to January 2022.

Removed in the revised version: To precisely elucidate the role of B cell subsets in anti-HTNV infection, we classified B cells into eight subsets, including antibody-secreting memory B cells (ASM), double-negative B cells (DN), intermediate memory B cells (IM), marginal zone-like B cells (MZB), naive B cells (naive B), plasmablasts (PB), quiescent resting memory B cells (RM), and exhausted tissue-like memory B cells (TLM) (Figure 2A and B, and Figure 2—figure supplement 1F). We discovered that significant expansion of ASM, PB, and RM cell populations in acute HFRS patients compared to healthy controls, accompanied by contraction of DN, IM, and naive B cell compartments post-infection (Figure 2C). In healthy controls, naive B cells (42.87%), DN B cells (26.65%), and TLM B cells (16.13%) constituted the majority of the B cell population. However, ASM (30.96%), PB cells (11.12%) and RM B cells (2.88%), became the dominant subsets after HTNV infection (Figure 2D and Supplementary file 1).

Immunophenotypic remodeling in the B-cell subsets during acute HFRS.

(A) t-distributed stochastic neighbor embedding (t-SNE) plot showing antibody secreting memory B cells (ASM), double negative (DN) B cells, intermediate memory cells (IM), marginal zone-like cells (MZB), naive B cells, plasmablasts (PB), quiescent resting memory cells (RM), and exhausted tissue-like memory B cells (TLM) of PBMCs identified using an integrated and classification analysis. (B) t-SNE projection of canonical markers, including CD19, CD27, CD38, IgD, and IgM. (C) Proportions of the eight B cell subsets, colored by the healthy group (green) and HTNV groups (red). Boxplot features: minimum box, 25th percentile; center, median; maximum box, 75th percentile. (D) Frequency of the eight B cell subsets in between the healthy group (n = 8) and HTNV groups (n = 15). (E) FACS gating strategy for the measurement of the B-cell subsets: activated memory B cells (CD21− CD27+, AM), RM B cells (CD21+ CD27+), IM B cells (CD21+ CD27−), TLM B cells (CD21− CD27−), naive B cells (CD27− IgD+), MZB B cells (CD27+ IgD+), ASM B cells (CD27+ IgD−), DN B cells (CD27− IgD−), PB (CD38+ CD27+), and class-switched memory B cells (CD38− CD27+, CSM), respectively. (F) The proportion of ASM, CSM, naive B, and PB cells in CD19+ B cells in the acute and convalescent HFRS patients. Data are presented as mean ± SEM in panels C and F. *p<0.05 (Student’s t-test).

Added in the revised version: To precisely elucidate the role of B cell subsets in anti-HTNV infection, we classified B cells into eight subsets, including antibody-secreting memory B cells (ASM), double-negative B cells (DN), intermediate memory B cells (IM), marginal zone-like B cells (MZB), naive B cells (naive B), plasmablasts (PB), quiescent resting memory B cells (RM), and exhausted tissue-like memory B cells (TLM) (Figure 2A and B, and Figure 2—figure supplement 1F). We discovered that significant expansion of ASM, PB, and RM cell populations in acute HFRS patients compared to healthy controls, accompanied by contraction of DN, IM, and naive B cell compartments post-infection (Figure 2C). In healthy controls, naive B cells (42.87%), DN B cells (26.65%), and TLM B cells (16.13%) constituted the majority of the B cell population. However, ASM (30.96%), PB cells (11.12%) and RM B cells (2.88%), became the dominant subsets after HTNV infection (Figure 2D and Supplementary file 1).

Removed in the revised version: Single-cell transcriptomes of PBMCs from patients with HFRS.

Removed in the revised version: (A) Uniform manifold approximation and projection (UMAP) presentation of the major peripheral immune cell types in the PBMCs from 15 HFRS patients and 8 healthy controls. (B) Dot plot shows the expression levels of canonical marker-related genes in the different subpopulations. (C) UMAP presentation of the major peripheral immune cell types among HFRS and healthy control groups. (D) Proportions of the eight cell subsets, colored by the healthy group (green) and HTNV groups (red). Boxplot features: minimum box, 25th percentile; center, median; maximum box, 75th percentile. (E) Frequency of the eight cell subsets in between the healthy group (n = 8) and the HTNV group (n = 15). (F) The UMAP plot shows the differentiation trajectories of different cell types calculated by pseudotime analysis.

Removed in the revised version: Dynamic analysis of the B-cell subsets in HFRS patients.

Removed in the revised version: (A) Expression of MZB in CD19+ B cells. (B) Expression of DN in CD19+ B cells. (C) Expression of TLM in CD19+ B cells. (D) Expression of AM in CD19+ B cells. (E) Expression of RM in CD19+ B cells. (F) Expression of IM in CD19+ B cells.

To validate the accuracy of our single-cell subpopulation analysis, we employed flow cytometry to quantify the B cell subsets in acute and convalescent HFRS patients (Figure 2E). The results demonstrated that ASM (39.48% vs. 24.57%, p<0.001) and PB (23.97% vs. 2.61%, p <0.001) subsets were significantly amplified in the acute phase compared to the convalescent phase (Figure 2F). Conversely, the number of naive B cells (32.78% vs. 48.92%, p=0.003) and CSM cells (18.65% vs. 27.15%, p=0.008) recovered rapidly in the convalescent stage (Figure 2F), consistent with the observations at the transcriptome level. No significant differences were observed for other B cell subtypes between the acute and convalescent stages (Figure 2—figure supplement 2). These results demonstrated that several B cell subpopulations were effectively activated following HTNV infection.

Dynamic characteristics of the IgG N-glycome in HFRS patients

HFRS patients were further stratified into three age groups (≤44, 45-59, and ≥60 years). IgG glycosylation was analyzed by ultra-performance liquid chromatography. As shown in Supplementary file 2, no significant differences in IgG glycosylation were observed between age groups, except for galactosylation (p=0.001). Similarly, no sex-based differences in IgG-Fc glycosylation were noted (Supplementary file 3). However, the levels of IgG with bisecting GlcNAc (15.72% vs. 14.51%, p<0.001), galactosylated IgG (74.26% vs. 71.15%, p<0.001), and sialylated IgG (22.44% vs. 21.62%, p<0.001) were significantly higher in the acute phase compared to the convalescent phase. Conversely, fucosylated IgG levels (94.29% vs. 94.89%, p<0.001) were lower in the acute phase (Supplementary file 4), indicating that IgG N-glycosylation plays a critical role in the anti-infection process. Multivariate linear regression was employed to mitigate potential confounding by genetic and environmental factors in the glycomics analysis. While no significant associations were observed for most glycan models (fucosylation, p=0.526; bisecting GlcNAc, p=0.069; and sialylation, p=0.058), we discovered sex showed a potentially influential effect on galactosylation (p=0.001) (Supplementary files 5-8). These results suggest that while most glycan features appear unaffected by the examined covariates, galactosylation may be subject to sex-specific biological regulation.

Among the 130 paired blood samples from 65 patients, 38 (29.2%) were negative for HTNV nucleocapsid protein-specific IgG antibodies, and 92 (70.8%) were seropositive. For the 24 patients with both seronegative and seropositive samples, levels of bisecting GlcNAc (16.16% vs. 14.54%, p<0.001), galactosylation (74.12% vs. 70.94%, p=0.006), and sialylation (22.92% vs. 21.80%, p=0.008) were higher in the seronegative period, whereas that of fucosylation was lower (94.36% vs. 94.74%, p= 0.024) in this period (Figure 3A). Notably, bisecting GlcNAc and sialylated IgG levels decreased as HTNV nucleocapsid protein (NP)-specific antibody titers increased (Figure 3B), highlighting a potential link between antibody levels and IgG-Fc glycosylation.

Glycosylation modification of antibodies associated with HTNV infection.

(A) Changes of different glycosylation types in both HTNV NP-specific IgG negative and positive plasma from 24 HFRS patients. (B) Differential IgG glycosylation patterns across antibody titer levels quantified by ELISA. (C) Comparison of different glycosylation levels before and after the fourfold increase in the IgG antibody titers. *p<0.05, **p<0.01, ***p<0.001. For paired samples, Wilcoxon Signed Rank Test was used to assess the difference.

The IgG N-glycome in HFRS patients with seroconversion

A four-fold or greater increase in HTNV NP-specific antibody titers usually indicates a protective humoral immune response during the acute phase (Onwuchekwa et al., 2023). We observed a significant increase in the Fc fucosylation (94.18% vs. 94.48%, p=0.044) and a decrease in bisecting GlcNAc (16.41% vs. 14.71%, p=0.001) post the four-fold antibody titer increase compared to the baseline. However, no significant differences were observed in galactosylation and sialylation (Figure 3C), suggesting that fucosylation and bisecting GlcNAc may play key roles in anti-HTNV infection and recovery.

The potential source of the IgG N-glycome during HTNV infection

Positive correlations were observed between the ASM subsets and both galactosylation (p=0.017, rs=0.418) and sialylation (p=0.008, rs=0.458) in the antibody Fc region, as well as between the PB subsets and sialylation (p=0.036, rs=0.372) (Figure 4A-C). Therefore, we speculated that galactosylated antibodies may primarily originate from the ASM subsets, while sialylated antibodies may derive from the ASM and PB subsets.

Glycosylation modifications of antibodies primarily deprived from antibody-secreting and plasmablast subpopulations.

(A) The correlation between the proportion of PB cells and the sialylation level. (B) The correlation between the proportion of ASM cells and the sialylation level. (C) The correlation between the proportion of ASM cells and the galactosylation level. (D) Dot plot shows the expression levels of glycosylation-related genes in the eight B cell groups. The pink color represents mannose related genes, the green color represents N-glycosylation-related genes, the orange color represents galactosylation-related genes, the purple represents sialylation-related genes, and the red color represents fucosylation related genes. (E) Enriched pathways in the plasmablast and ASM subsets by GO enrichment analysis of the DEGs. (F) The GSEA map presents the enrichment score of GRGs in N-Glycan biosynthesis.

Removed in the revised version: The differentially expressed genes and their functional changes in the B cell subsets post HTNV infection.

Removed in the revised version: (A) The volcano plot shows the differentially expressed genes of different B cell subsets in acute HFRS patients. The red font represents the upregulation of the B cell subpopulations, and the blue font represents the downregulation of the B cell subpopulations. (B and C) The volcano plot displays the differentially expressed genes in upregulation B cell subpopulations including antibody secreting memory B cells, plasmablasts, and quiescent resting memory B cells (B), and downregulation B cell subpopulations including naive B cells, double negative B cells, and intermediate memory B cells (C).

To investigate the process of glycosylated antibody production secreted by the ASM and PB subsets, we identified differentially expressed genes (DEGs) among the eight B cell subpopulations. Notably, the most DEGs were observed in the ASM (470), naive (312), and TLM (333) B subpopulations, respectively (Figure 4—figure supplement 1A and Supplementary file 9). Specifically, among all the upregulated genes, the expressions of glycosylation-related, mitochondrial respiratory-related, and peptide chain elongation-related genes among the ASM, PB, and RM populations were elevated after HTNV infection, including RPN2, MT-ND6, and RPS26 (Figure 4—figure supplement 1B and Supplementary file 10). On the contrary, the upregulated genes in the naive, DN, and IM populations were related to anti-inflammatory and immunosuppressive effects, inhibition of the antioxidative function, response to environmental stress, and B-cell growth and development, such as TSC22D3, TXNIP, DUSP1, and KLF6, whereas the downregulated genes played an important role in the inflammatory response, cell motility, and antigen presentation, including LTB, ACTG1, and HLA-DRB5 (Figure 4—figure supplement 1C and Supplementary file 11).

Unexpectedly, we observed significant changes in the expression of glycosylation-related genes (GRGs) in the ASM, PB, and RM subclusters during HTNV infection compared to healthy controls. RPN1 and RPN2 were prominently upregulated, which encoded subunits of the oligosaccharyltransferase complex (Figure 4D). This implied that the process of glycosylation during HTNV infection was rapidly activated among the ASM and PB subclusters. In the ASM subcluster, the fold change of GRGs from all four glycosylation modification-related genes, aside from RPN1 and RPN2, was similar, suggesting that ASM may undergo complex and diverse glycosylation modification processes.

To enhance our understanding of the transcriptomic characteristics of GRG-high-expressing PB and ASM subclusters, we conducted a comprehensive pathway enrichment analysis and identified significant enrichment of pathways related to glycosylation modifications, including the endoplasmic reticulum protein-containing complex, rough endoplasmic reticulum, and protein N-linked glycosylation via asparagine (Figure 4E). Additionally, gene sets associated with glycosylation modifications showed positive regulation of these pathways (p=0.002, Figure 4F). These results indicated that glycosylation modifications in the ASM, PB, and RM subpopulations were dynamically regulated during the acute phase of HFRS, likely to satisfy the amplified anti-HTNV antibody requirement and accelerate viral clearance.

Discussion

IgG glycosylation is a key regulator of inflammatory responses for infectious diseases, acting as a molecular switch between anti-inflammatory and pro-inflammatory effects upon antigenic challenge. One of the most significant findings of the present study is the elevation in IgG fucosylation and reduction in bisecting GlcNAc levels at the time of a four-fold increases in the HTNV NP-specific antibody titers. Afucosylation is known to enhance antibody-dependent cellular cytotoxicity (ADCC) by increasing the affinity of IgG-Fc for the FcγRIIIa receptor on NK cells, leading to heightened production of inflammatory cytokines such as IL-1β, IL-6, IFN-γ, and TNF-α (Larsen et al., 2021; Lu, Suscovich, Fortune, & Alter, 2018; Wieczorek, Braicu, Oliveira-Ferrer, Sehouli, & Blanchard, 2020). Increase in bisecting GlcNAc linked to Fc induces enhanced affinity for the FcγIIIa receptor and enhanced ADCC function of the antibodies (Giron et al., 2020; Hou et al., 2021; Irvine & Alter, 2020). However, rather than directly influencing the Fc receptor binding, this modification enhances antibody functionality through afucosylation, which subsequently increases the binding affinity to the FcgRIIIa receptor (Ackerman et al., 2018; Jennewein & Alter, 2017). Our findings provide evidence that the dynamic alterations in glycosylation patterns, characterized by increased fucosylation and decreased bisecting GlcNAc levels during viral infection, are involved in disease pathogenesis and clinical recovery, suggesting their dual utility as both diagnostic biomarkers and promising therapeutic targets.

We also observed elevated galactosylation levels in the acute phase compared to the convalescent phase, a phenomenon also reported in SARS-CoV-2 infections but not in chronic diseases and cancers (Hou et al., 2021; Liu et al., 2018). Increased galactosylation may expand the distance between CH2 domains, exposing key amino acids for FcγRIIIa binding and enhancing ADCC functions (Houde, Peng, Berkowitz, & Engen, 2010). Simultaneously, the increase in the ASM subset, which positively correlated with galactosylation, may reflect polyclonal activation of memory B cells driven by pattern recognition receptors and cytokines, leading to the production of galactosylated antibodies. Similarly, sialylation levels were higher in the acute patients, potentially linked to the expansion of the ASM and PB subsets. In this study, the PB subset was significantly elevated in the acute phase and positively correlated with sialylation, suggesting that rapidly proliferating plasmablasts may contribute to the secretion of sialylated antibodies; however, the mechanisms behind this require further investigation.

Antibody-related glycogenes are significantly activated following Hantaan virus infection. We noted that ribophorin I and II (RPN1 and RPN2) were significantly upregulated in the ASM/IM/PB/RM subsets after Hantaan virus infection, which linked the high mannose oligosaccharides with asparagine residues found in the Asn-X-Ser/Thr consensus motif (Hwang et al., 2025). We speculate that they continuously attach the synthesized glycan chains to the constant region of antibodies during antibody synthesis. Similarly, fucosyltransferase 8 (FUT8) in the ASM subset, catalyzing the alpha1-2, alpha1-3, and alpha1-4 fucose addition (Wang & Ravetch, 2019; Yang et al., 2015), was downregulated in the mRNA translation, and the levels of fucosylated antibodies were naturally lower in the acute HFRS patients. Meanwhile, the beta-1,4-galactosyltransferase (beta4GalT) gene expression was significantly elevated in the ASM subpopulation during the acute phase, which also correlated with increased levels of galactosylated antibodies in serum (Wang & Ravetch, 2019). However, we did not observe significant upward changes in sialyltransferase mRNA expression in the acute HFRS patients, similar with the finding from severe COVID-19 cohorts (Haslund-Gourley et al., 2024). The neuraminidase 1 (NEU1) gene is strikingly upregulated and may potentially explain the decreased sialylation on the secreted HTNV-specific IgG antibodies during convalescence. Overall, the glycosylation of immunoglobulin G is regulated by a large network of B-cell glycogenes during HTNV infection.

Despite establishing a well-characterized patient cohort and performing systematic IgG glycosylation profiling based on HTNV NP antibody status, this study has several noteworthy limitations. Most notably, while preliminary comparisons suggested similar patterns between virus-specific and total IgG N-glycome, our total plasma IgG analysis may have introduced confounding factors in the observed associations. This methodological constraint could potentially affect the interpretation of certain disease-specific glycosylation signatures.

Our findings demonstrate that the IgG-Fc glycosylation plays a pivotal role in HTNV infection and rapid recovery. Fucosylated oligosaccharides and the absence of bisecting GlcNAc are associated with favorable clinical outcomes. Importantly, specific B cell subsets, particularly ASMs and PBs, upregulated the secretion of galactosylated and sialylated antibodies, respectively. These findings would deepen our understanding of antibody glycosylation in antivirus immunity and provide a foundation for future studies aimed at optimizing glycoengineered therapeutic antibodies.

Materials and Methods

Study participants

Clinical specimens were collected from HFRS patients who were hospitalized in Baoji Central Hospital between October 2019 and January 2022. Patients were categorized into four clinical subtypes (mild, moderate, severe, and critical) based on the diagnostic criteria for HFRS issued by the Ministry of Health (Ma et al., 2015). This study was approved by the ethics committee of the Shandong First Medical University & Shandong Academy of Medical Sciences (R201937). Written informed consent was obtained from each participant or their guardians.

The clinical course of HFRS is grouped into acute (febrile, hypotensive, and oliguric stages) and convalescent (diuretic and convalescent stages) phases. The acute phase was defined as within 12 days of illness onset, and the convalescent phase was defined as a period of illness lasting 13 days or longer (Tang et al., 2019; Zhang et al., 2022). The earliest sample was selected if there were multiple blood samples available in the acute phase and the last available sample before discharge was selected if there were multiple blood samples in the convalescent phase.

Single-Cell RNA Sequencing and public database

Flow cytometry antibody staining was conducted as previously described (Chakraborty et al., 2022). Dead cell and red blood cell were removed using the BD FACSAria™ III cell sorter, and peripheral blood mononuclear cells (PBMCs) were collected for single-cell sequencing. The library construction was conducted as previously described (Jin et al., 2024).

We obtained single-cell transcriptome sequencing data from the Gene Expression Omnibus (GEO) database (GSE161354), containing six HFRS patients and two healthy volunteers, as well as healthy control samples from the Genome Sequence Archive (GSA) database (HRA000203), matched with patient demographic information.

Single-Cell Data Analysis

The FastQC software was used to evaluate the data obtained to ensure the quality of the raw sequencing data. The raw data were mapped to the human reference genome (GRCh38, https://cf.10xgenomics.com/supp/cell-exp/refdata-gex-GRCh38-2020-A.tar. gz) using Cell Ranger. The Seurat R package (Version 4.4.0) was utilized for the merging and clustering of single-cell data. The Harmony R package and the anchor module of Seurat were used to remove batch effects between samples and groups for cell clustering. t-distributed Stochastic Neighbor Embedding (t-SNE) and Uniform Manifold Approximation and Projection (UMAP) were employed for dimensionality reduction and visualization of individual cells.

Cell Type Annotation

UMAP and t-SNE were used to reduce the dimension of all cells and to cluster them in a two-dimensional space based on shared features. The Single R (Version 2.2.0) package was used to infer independently the cell origin of each single cell based on a reference transcriptome dataset of pure cell types, for unbiased cell type identification from single-cell RNA sequencing data. Subsequently, specific highly-expressed genes were combined with those derived from review of the literature for manual annotation to ultimately determine cell types. Classic biomarkers for specific cell types were used to identify cells in different clusters.

Differential Gene Identification and Functional Analysis

The FindMarkers() function in the Seurat package was used to identify DEGs between different cell groups, with criteria of |log2FC| > 0.25 and p-value<0.05, pct.1>0.25. The ClusterProfiler R package was utilized for GO, KEGG, and GSEA enrichment analyses, and the ggplot2 and enrichplot R packages were used for visualization.

Measurement of HTNV specific antibodies

Levels of serum IgM and IgG antibodies to HTNV NP were assessed through the enzyme-linked immunosorbent assay (ELISA) kit (Wantai BioPharm, Beijing, China) in accordance with the manufacturer’s instructions.

Analysis of IgG glycans

The process of isolating, labelling, purifying, and analyzing plasma IgG was performed using previously established methods (Hou et al., 2019; Pengcheng Liu, 2023). Briefly, the diluted plasma samples were transferred onto a 96-well protein G monolithic plate (BIA Separations, Slovenia) for the isolation of IgG. The isolated IgG was eluted with 1 mL of 0.1 M formic acid and was immediately neutralized with 170 µL of 1M ammonium bicarbonate.

The released N-glycans were labelled with 2-aminobenzamide (2-AB) and were then purified from a mixture of 100% acetonitrile and ultrapure water in a 1:1 ratio (v/v). This was then analyzed by hydrophilic interaction liquid chromatography using ultra-performance liquid chromatography (HILIC-UPLC; Walters Corporation, Milford, MA) (Hou et al., 2019). As previously reported, the chromatograms were separated into 24 IgG glycan peaks (GPs) (Menni et al., 2018).

Earlier version: Our study incorporated both biological and technical replicates to ensure a robust glycomic profiling analysis. Specifically, we analyzed paired acute/convalescent-phase samples from 65 confirmed HFRS patients to assess inter-individual biological variability, while technical reproducibility was validated through comparison with standard chromatographic peak plots (Vučković et al., 2016). This dual-replicate strategy enabled a comprehensive evaluation of both biological heterogeneity and assay precision.

Revised version: Our study incorporated both biological and technical replicates to ensure a robust glycomic profiling analysis. Specifically, we analyzed paired acute/convalescent-phase samples from 65 confirmed HFRS patients to assess inter-individual biological variability, while technical reproducibility was validated through comparison with standard chromatographic peak plots (Vučković et al., 2016). This dual-replicate strategy enabled a comprehensive evaluation of both biological heterogeneity and assay precision, and the coefficient of variation for samples were below 16%.

Virus-specific B-cell detection

Virus-specific B-cell against HTNV were performed as previously described (Liechti & Roederer, 2019; Song et al., 2018).

Statistical analysis

Categorical variables were compared using chi-square tests. Normality of numeric variables was assessed with the Kolmogorov-Smirnov test; non-normally distributed continuous variables were analyzed using non-parametric tests. The Benjamini - Hochberg (BH) method was used to adjust the raw p-values from DEG analysis, controlling the false discovery rate (FDR). Statistical analyses were performed with R software version 4.1.1 (R Core Team, New Zealand) and GraphPad Prism 8 (GraphPad Software, San Diego, CA), with a significance level of α = 0.05 (two-tailed).